# Problem 3: Sorting

You have been provided with an unsorted list of numbers. Write a python function to put the numbers in ascending order. You may not use the built-in sort command in python – the point of the exercise is to think about how to design an algorithm to solve this problem yourself. Your code should consist of a script in which you load the data, call a function (discussed next), and print the sorted list and a function that takes as argument the unsorted list and returns the sorted list. Both your script and your function should be adequately commented (see the style guide) so that your TA and instructor can see very easily what your program is doing.

*Note: This last exercise is a solved problem in the sense that many extremely efficient sorting algorithms exist. The point of the exercise is less to produce optimally efficient code and more to develop your own algorithmic thinking and python programming skills. Copying code from the internet will result in a grade of ‘0’ for the whole lab. Paraphrasing code from the internet might get you a good grade on the lab, but will hurt you on the midterm.*

Describe in words how your sorting algorithm works, step-by-step. This should be described with a level of detail that would allow a classmate to reproduce your code without seeing it.

<!-- grade_id: description, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

Now add code to the function my_sort below to implement your sorting algorithm:

<!-- grade_id: sort_algo, points: 0, grade: false, solution: true -->
```python
def my_sort(unsorted_list):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test your function by generating and sorting a list of random numbers:
x1 = np.random.randint(0,high=10000,size=(10))
print("x1=",x1)
print("my_sort(x1)=",my_sort(x1))
```

<!-- grade_id: sort_tests, points: 5, grade: true, solution: false -->
```python
import numpy as np

x1 = np.random.randint(0,high=10000,size=(100) )
x2 = np.random.randint(0,high=10000,size=(1000))
assert( (np.sort(x) == my_sort(x)).all()   )
assert( (np.sort(x2) == my_sort(x2)).all() )
```