## Problem 2
You have been provided with a set of x,y coordinates describing the location of the center line of a river. The sinuosity of a river is defined as:

$S = \frac{\text{actual path length}}{\text{shortest path length}}$

Sinuosity is one of the most important parameters in fluvial geomorphology. Your task is to read in the data in ‘river_centerline.mat’, which contains two variables e_m and n_m that contain coordinates of the channel centerline in meters east and north of a reference point. Write a program that calculates the sinuosity for this reach of the river in question.

Write code to load the river centerline data from river_centerline.txt. Use the `loadtxt` command from `numpy`. Store the `x` or East coordinate in a variable called e_m and the `y` or North coordinate in a variable called n_m.


<!-- grade_id: loadtext, points: 0, grade: false, solution: true -->
```python
"""Write code to load the river centerline data from river_centerline.txt. Use the `loadtxt` command from numpy"""
# YOUR CODE HERE
raise NotImplementedError()
```

```python
"""Check the shape of the variables that you just loaded"""
np.shape(e_m)
```

<!-- grade_id: centerline, points: 2, grade: true, solution: false -->
```python
"""check for existence and correct shape of e_m and n_m"""
import numpy as np
if 'e_m' in locals() :
    assert( np.shape(e_m) == (741,) )
else:
    raise(NameError)
if 'n_m' in locals() :
    assert( np.shape(n_m) == (741,) )
else:
    raise(NameError)
```

Write a function called dist(x1,y1,x2,y2) that calculates and returns the straight-line distance between the points with coordinates (x1,y1) and (x2,y2)

<!-- grade_id: distance, points: 0, grade: false, solution: true -->
```python
def dist(x1,y1,x2,y2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_dist, points: 3, grade: true, solution: false -->
```python
"""Test the distance function"""
assert( dist(0,0,2,0) == 2)
assert( dist(-1,-1,-1,0) == 1)
assert( np.abs(dist(-1,-1,0,0) - np.sqrt(2)) <= 1E-16 )
```

Now, write a function called sinuosity that takes two arguments: x and y. calculate and return the sinuosity. Use the dist function that you wrote previously.

<!-- grade_id: sinuosity, points: 0, grade: false, solution: true -->
```python
def sinuosity(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sinuosity(e_m,n_m)
```

<!-- grade_id: test_sinuosity, points: 5, grade: true, solution: false -->
```python
"""Verify the correctness of the sinuosity result (within one percent)"""
assert( np.abs( sinuosity(e_m,n_m)-1.7367516388813369 ) <= 0.017 )
```

<!-- grade_id: test_usedist, points: 3, grade: true, solution: false -->
```python
"""Verify that sinuosity used the dist function"""
orig_dist = dist
del dist
try:
    sinuosity([0,1],[0,1])
except NameError:
    pass
else:
    raise AssertionError("sinuosity() does not use dist()")
finally:
    dist = orig_dist
```

Finally, use matplotlib to make a plot of the river centerline. Your plot should:

1. Have equal axis aspect ratio (i.e. 1 m on the x-axis should be the same size as 1 m on the y-axis)
2. Have labeled axes
3. Show the straight-line path from the first data point to the last data point with a dashed line in a contrasting color.
4. Include a legend indicating the 'centerline' and 'shortest path'
5. Have a descriptive title.

<!-- grade_id: map_plot, points: 5, grade: true, solution: true -->
```python
"""Put your plotting code here:"""
# YOUR CODE HERE
raise NotImplementedError()
```