# Problem 1:

You will write a script that calculates and displays the first 20 Fibonacci numbers. Recall that the first Fibonacci number $F_1=0$ and the second Fibonacci number $F_2=1$. The rest of the Fibonacci numbers can be generated through iteration using the recurrence relation $F_n=F_{n-1}+F_{n-2}$.

Write a function called next_fib that takes two arguments and returns the next Fibonacci number in the sequence:


<!-- grade_id: nextfib, points: 0, grade: false, solution: true -->
```python
def next_fib(a,b):
    # YOUR CODE HERE
    raise NotImplementedError()
    raise(NotImplementedError)
```

Test the function that you just implemented. `next_fib(0,1)` should return $0+1=1$.

<!-- grade_id: test_nextfib, points: 1, grade: true, solution: false -->
```python
assert(next_fib(0,1)==1)
assert(next_fib(1,1)==2)
assert(next_fib(1,2)==3)
```

Now write a function called fibonacci that takes one argument `n` and returns the list of the first `n` Fibonacci numbers. If $n<1$, your code should raise a `ValueError`. Your code must achieve this by calling the function `next_fib` that you defined previously.

<!-- grade_id: fibonacci, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Test the function that you wrote. The first 5 fibonacci numbers are $[0,1,1,2,3]$. Verify that this is the answer that your code produces:

```python
fibonacci(5)
```

<!-- grade_id: test_fib, points: 5, grade: true, solution: false -->
```python
"""Test the fibonacci function for various use cases"""
assert( fibonacci(1) == [0] )
assert( fibonacci(2) == [0,1] )
assert( fibonacci(5) == [0,1,1,2,3] )
assert( fibonacci(10) == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34] )
```

<!-- grade_id: test_rely, points: 3, grade: true, solution: false -->
```python
"""Verify that fibonacci relies on next_fib"""
orig_nextfib = next_fib
del next_fib
try:
    fibonacci(5)
except NameError:
    pass
else:
    raise AssertionError("fibonacci does not use next_fib")
finally:
    next_fib = orig_nextfib
```