# G326: Lab 3
This week we will use the computer as tool to solve a differential equation. The differential equation that you will solve describes the time evolution of many zero-dimensional (box) models. The form of the solution is by now very familiar and leads to exponential behavior in time. In general, when an analytic solution exists, it is preferable to use this solution. However, because this may be the first time that you've solved equations numerically, I have chosen a problem that has an analytic solution so that you can *verify* that your numerical solution is correct. The numerical solutions that you obtain should all be very similar to the exact solution (within a few percent, usually much closer).

Recall the general form of the decay equation:

$$\begin{equation}
\frac{d u}{dt} = -\lambda u                
\end{equation}$$

subject to the Initial Condition (IC):

$$u(0) = u_0$$.

Here, $u(t)$ is the amount of a radioisotope. This equation has a simple analytic solution for $u(t)$ with which you are familiar. In general, when we obtain an analytic solution to a differential equation, we obtain an expression that allows us to calculate the value of $u(t)$ for every time $t$. On the other hand, when we obtain a numerical solution to this equation, we will calculate approximate values of the solution at specific, discrete, times. For now, we will assume that these times are uniformly spaced, with spacing $k$. We refer to this spacing $k$ as the time step. We use a superscript $n$ to denote the time corresponding to the $n$-th time. We can state mathematically the relationship between the $n$-th time and the preceding time:

$$\begin{equation}
t^n = t^{n-1}+k
\end{equation}$$

**Note that here the superscript $n$ is NOT an exponent.** It is an index referring to the $n$-th time at which we obtain an approximation to the solution. The reason that we use this seemingly confusing notation is that we will later use subscripts to denote spatial indices.


We will refer to the approximation to the function $u(t)$ at time $t=t^n$ as $u(t^n)=u^n$. Again, $n$ is **not** an exponent - it is a superscript identifying the time at which a solution is obtained.
Now that we have *discretized* time and introduced notation for the values of the dependent and independent variables at each time step, we need to discuss how we might go about approximating the derivative $\frac{d u}{dt}$ numerically. We will consider four choices today and express them in terms of the general first-order ordinary differential equation:

$$\begin{equation}
\frac{d u}{dt} = f(u(t),t)
\end{equation}$$

In all cases, we obtain a numerical approximation by introducing a discrete approximation to the derivative $\frac{du}{dt}$ at time $t^n$. The first choice is referred to as \textit{forward Euler} or sometimes \textit{Euler's method} and involves the approximation:

$$\begin{equation}
\frac{d u}{dt}\bigg|_{t=t^{n}} = \frac{u^{n+1}-u^{n}}{k} 
\end{equation}$$

$$\begin{equation}
\frac{u^{n+1}-u^{n}}{k} = f(u^{n},t^{n})
\end{equation}$$

**Important: In expressions involving superscripts that refer to a particular instant in time (e.g. $t^n$), we can always replace $n$ with any other expression. For instance, we can substitute $m=n+1$ in the previoux equation to obtain:**

$$\begin{equation}
\frac{u^{m}-u^{m-1}}{k} = f(u^{m-1},t^{m-1})
\end{equation}$$

Note that here, when calculating the value of $u(t)$ at time $t^{n}$, the function $f$ is evaluated at the previous time $t^{n-1}$. Because this method involves only the value of the solution at time $t^{n-1}$, it is called an *explicit* method.


**Part 1:** What is the analytic solution to equation 1? *(Hint: This is a warm-up problem and should be pretty easy after last week's lab!)*



<!-- grade_id: q1_ans, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 2:**
1. Create a time vector called `t`. It is good programming practice in general to first create variables representing the initial time, the final time, and the number of time steps. Then, if you need to refer to these values elsewhere in your program, you can use the variable names rather than typing in a numerical constant. If you decide, for instance, that you need to revise your code and take smaller time steps (this is a hint of things to come...), you can do so with a one-line change instead of changing every instance of the number of time steps. 

2. Define nt as the number of time steps to be taken, tmin as the starting time, and tmax as the ending time. For now, you should begin at time t=0, end at time t = 4.5 Gyr, and take 20 time steps. 

3. Calculate k (the time step) and assign it to a variable called k. We will use the value λ = 0.6/Gyr throughout this assignment. 

4. In this first code block, it is good practice to define all of the parameters that you will need to refer to in the code. Define a variable called `lam` and refer to this variable name rather than entering numerical constants. Also define a variable called `u0` that holds the initial value, $u_0$. Assume that the initial value $u_0=3\times 10^5$.


<!-- grade_id: q2_response, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q2_t1, points: 3, grade: true, solution: false -->
```python
# Check that t is the right size, with correct first and last values
assert( nt == 20 )
assert( len(t) == nt )
assert( t[0] == 0. )
assert( t[-1] == 4.5)
```

<!-- grade_id: q2_t2, points: 2, grade: true, solution: false -->
```python
# Check that the time step and the time vector are set consistently
assert( np.abs( (t[1]-t[0]) - k < 1e-15 )  )
```

<!-- grade_id: q2_t3, points: 2, grade: true, solution: false -->
```python
# Check for correct values of lam and u0.
assert( lam==0.6 )
assert( u0== 3.0e5 )
```

**Part 3** Create a new vector called `uexact` that contains the exact values of u(t) calculated from your analytic solution from Part 1. The values in `uexact` here should correspond to the times listed in the time vector in Part 2. 

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q3_t2, points: 2, grade: true, solution: false -->
```python
assert( uexact.shape == t.shape )
```

**Part 4** Create a new vector called ‘ufe’ with the same dimensions as your time vector. This vector will hold your numerical solution to equation (1). Set the first element in ufe, which corresponds to time t=0, to the initial value of u.

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q4t1, points: 4, grade: true, solution: false -->
```python
# Note: it is really important that these tests pass - ask for help if your code is not working.
assert( ufe.shape == t.shape )
assert( ufe[0] == u0 )
```

**Part 5a** Derive a formula (and write it here) for $u^{n+1}$ in terms of $u^n$ using the forward Euler method described above.

<!-- grade_id: q5a1, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 5b** Use a ‘for’ loop to loop over the remaining time steps and populate ufe[1] through ufe[nt-1], where nt is the total number of time steps taken. 

<!-- grade_id: q5a2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Part 5c** Make a plot showing time on the horizontal axis and 'u' on the vertical axis. Plot your numerical solution `ufe` using a solid black line and plot the analytic solution `uexact` using a red dashed line. Add a legend labeling the exact and forward Euler solutions. *Hint: they should agree quite well*

<!-- grade_id: q5c1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q5ct1, points: 5, grade: true, solution: false -->
```python
# Check for good agreement between exact and numerical solutions
assert( np.linalg.norm(ufe-uexact)/np.linalg.norm(uexact) <= 0.20 )
assert( np.abs(ufe[-1]-16306.746517096166) < 160 )
```

# Part 6: Backwards Euler method
Our second choice of time integration scheme is called the Backwards Euler method:
$$\begin{equation}
\frac{d u}{dt}\bigg|_{t=t^{n}} = \frac{u^{n}-u^{n-1}}{k} 
\end{equation}$$

$$\begin{equation}
\frac{u^{n}-u^{n-1}}{k} = f(u^{n},t^{n})
\end{equation}$$
Note that here, we need to be able to solve the above equation for $u^{n}$, which may not be possible for some choices of $f$. Because this method requires knowledge of $f(u^{n},t^n)$ it is called an *implicit* method.

Derive a formula (and write it here) for $u^{n+1}$  in terms of $u^n$ using the Backwards Euler method: (show your work)

<!-- grade_id: q6a1, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

Now, create a new array called `ube` with the same dimensions as the time vector. Use a `for` loop to calculate an approximate solution using the Backwards Euler method:

<!-- grade_id: q6b1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q6bt2, points: 4, grade: true, solution: false -->
```python
# Check that the backwards Euler solution correct
assert(np.abs(ube[-1]-24027.2866541) < 240.0)
```

**Part 7: Runge-Kutta method**
A third scheme that we will consider belongs to a family of numerical integration schemes called Runge-Kutta methods. This one is referred to as a second-order Runge-Kutta scheme. This scheme involves taking a `half-step', then approximating the derivative at the midpoint between $t^{n-1}$ and $t^{n}$ and using this derivative to calculate the new value of $u^n$.
$$\begin{equation}
u^*=u^{n-1} + \frac{k}{2}f(u^{n-1},t^{n-1})
\end{equation}$$

$$\begin{equation}
u^{n}=u^{n-1} + k f(u^*,t^{n-1/2})
\end{equation}$$
Note that what we are doing in equation here is to approximate the value of $u$ at the midpoint between time $t^{n-1}$ and $t^{n}$ using the Forward Euler method. In the second step, we approximate the derivative $\frac{d u}{dt}$ by evaluating $f(t,u(t))$ using the value of $u$ calculated at $t^{n-1/2}=t^n-k/2$. This method is $explicit$ and hence does not require solving an equation involving $f$.

Write code to create a vector called `urk` with the same dimensions as the time vector and populate it a numerical solution obtained using the Runge-Kutta method:

<!-- grade_id: q7a1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q7at1, points: 4, grade: true, solution: false -->
```python
# Check that the backwards Euler solution correct
assert(np.abs(urk[-1])-20366.5295543 < 5.0)
```

**Part 8** Make a plot showing all four solutions (uexact, ufe, ube, and urk) with distinct line styles. Include a legend and, as always, label your axes appropriately.

<!-- grade_id: q7a1a, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Part 9** 
For each scheme (Forward Euler, Backward Euler, and Runge-Kutta), **calculate and print** the relative error. One measure of the relative error is given by the 2-norm:
$$\begin{equation}
\epsilon_2 = \frac{||u-u_{exact}||_2}{||u_{exact}||_2}
\end{equation}$$
Here, $||\cdot ||_2$ denotes the 2-norm. For a vector $\underline{v}$ of $n$ elements, the 2-norm is defined as
$$\begin{equation}
||\underline{v}||_2 = \sqrt{\sum_{i=1}^n v_i^2}
\end{equation}$$
Note that there is a built-in numpy function called `numpy.linalg.norm` that can calculate vector norms for you.

<!-- grade_id: q8a1, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Which scheme is most accurate when using 20 timesteps? Which is least accurate?

<!-- grade_id: q8q2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 10:**
Try changing the number of time steps nt=(10,100,1000,10000), and calculate the relative error (Part 9) for each choice of the time step. Make a **log-log** plot showing the size of the timestep taken on the horizontal axis and the relative error on the vertical axis. Plot the relative error for each scheme using a different choice of line color or symbol (i.e. you should have one line for the Runge-Kutta scheme, one line for the Forward Euler scheme, and one line for the Backward Euler Scheme). Note - you will have to either run your code above repeatedly with different values of `nt` (either in a separate notebook or this one) and make a data table containing the relative error values, OR you can write a more complicated `for` loop that iterates over all of the choices of time step and makes the plot, all in one go.

<!-- grade_id: q9a1, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Which numerical integration method is most accurate?

For which method(s) does the error decrease most quickly with decreasing step size?

For which method(s) does the error decrease least quickly with decreasing step size?

<!-- grade_id: q9p2a, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

# Bonus Problem (for extra credit):
Consider the ODE:

$$\begin{equation}
\frac{du}{dt} = \cos(2\pi t)
\end{equation}$$

With initial condition $u(0) = 1$. 

What is the analytic solution to this initial value problem?

<!-- grade_id: bonus2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Write a program to solve this ODE using the forward Euler method over the time interval $0\le t\le 10$. Explore choices of the timestep such that you take 10,100,1000, and 10000 time-steps between $t=0$ and $t=10$. For each case, plot the exact and numerical solutions using lines that can be distinguished.

<!-- grade_id: bonus4, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Describe qualitatively how the numerical and exact solutions differ when you take large timesteps. How small a timestep is necessary to obtain a relative error of less than $10^{-6}$?

<!-- grade_id: bonus6, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE