# Introduction


This week's lab focuses on steady-state solutions to the diffusion equation in one spatial dimension. Like the decay equation that you worked with last week, the diffusion equation is a differential equation, but unlike the decay equation, the diffusion equation is a second order partial differential equation (in general) that can be reduced to a second order ordinary differential equation in the special case of steady-state, one-dimensional diffusion. The equation that we'll be solving has the general form:

**Equation 1:**
$$
-\frac{\partial }{\partial x}\left(-k \frac{\partial u}{\partial x}\right) +Q=0  
$$

In the above equation, $u$ is the dependent variable representing the quantity that is diffusing (e.g. temperature, the concentration of a chemical compound, or hydraulic head), $x$ is a spatial coordinate, and $Q$ represents any sources or sinks of the diffusing quantity (for instance production of heat by radioactivity within a planet or creation/destruction of a chemical compound through reaction). Because the diffusion equation is a second-order differential equation, it requires *two* boundary conditions. Because there is no time-dependence (more on that next week), we do not need to specify initial conditions.

Like previous exercises, we will obtain a discrete solution to the equation at a series of points. However, rather than being points in time, these are now points in space. We'll find the solution at a number (`nx`) of points. We will use subscripts to refer to these points:
$$
\underline{x} = \left[x_0,x_1,\dots,x_{nx-1}\right]^T
$$
Note the presence of the transpose symbol $^T$ - in this week's exercise, $\underline{x}$ and $\underline{u}$ are column vectors.

![Figure 1](grid.png)

**Figure 1**: *Grid on which numerical solution will be obtained, showing points at which $T$ and $dq/dx$ are evaluated and points at which $q(x)$ is evaluated. The formulas give the coordinates $xc_i$ of the points at which $q(x)$ is evaluated (the midpoint of each interval between points $x_i$).*

Our goal is to obtain a solution $\underline{u}$ that is a vector containing the values ($u_i$) of our numerical solution to equation Equation 1 at each point $x_i$. First, recall that

$$
q(x) = -k\frac{\partial u}{\partial x}
$$

The first step in constructing our finite difference approximation to $u(x)$ is to construct a finite-difference approximation to $q(x)$ at the points $xc_{i-1}$ (the midpoint between $x_{i-1}$ and $x_i$) and $xc_{i}$ (the midpoint between $x_{i}$ and $x_{i+1}$). These points are shown graphically in Figure 1.

**Equation 2:**
$$
\begin{split}
q\bigg|_{x=xc_{i-1}} = -k \frac{u_i-u_{i-1}}{x_i-x_{i-1}}\\
q\bigg|_{x=xc_{i}} = -k \frac{u_{i+1}-u_{i}}{x_{i+1}-x_{i}}
\end{split}
$$

The second step in constructing our finite difference approximation to $u(x)$ is to write an approximation to $-\frac{\partial q(x)}{\partial x}$ at the point $x_i$:

**Equation 3**
$$
\left. \frac{\partial q}{\partial x}\right\rvert_{x=x_i} = \frac{q( xc_i ) - q( xc_{i-1} )}{xc_i-xc_{i-1}}
$$

Finally, we can substitute equation Equation 2 into Equation 3 to obtain an expression involving $u$:
$$
\left. -\frac{\partial}{\partial x}\left(-k \frac{\partial u}{\partial x} \right) \right\rvert_{x=x_i} = \frac{-1}{xc_i-xc_{i-1}}\left( -k \frac{u_{i+1}-u_{i}}{x_{i+1}-x_{i}} - \left(-k \frac{u_i-u_{i-1}}{x_i-x_{i-1}}\right)\right)
$$
The above expression can be used to obtain a numerical approximation to the solution $u(x)$ on a grid with arbitrary spacing. For the special case where the grid spacing is uniform and equal to $h$ (that is, $x_i=x_{i-1} + \Delta x$), we can simplify the previous equation to:

$$
\left. -\frac{\partial}{\partial x}\left(-k \frac{\partial u}{\partial x} \right) \right\rvert_{x=x_i} = \frac{-1}{\Delta x}\left( -k \frac{u_{i+1}-u_{i}}{\Delta x} - \left(-k \frac{u_i-u_{i-1}}{\Delta x}\right)\right)=\frac{k}{(\Delta x)^2}\left( u_{i+1} - 2 u_i + u_{i-1} \right)
$$

Thus, our approximation to equation 1 for the special case with uniform grid spacing becomes:

$$
\frac{k}{(\Delta x)^2}\left(u_{i-1}-2 u_{i}+u_{i+1}\right) + Q_i = 0
$$

Note that this equation involves the value of the solution at the $i$-th point as well as its immediate neighbors to the left and right. If there are $nx$ grid points, the above finite difference approximation leads to $nx$ equations for $u_i$. For instance, for $i=2$, we have:

$$
\frac{k}{(\Delta x)^2}\left(u_{1}-2 u_{2}+u_{3}\right) + Q_i = 0
$$

and for $i=3$, we have:

$$
\frac{k}{(\Delta x)^2}\left(u_{2}-2 u_{3}+u_{4}\right) + Q_i = 0
%\frac{u_{2}-2 u_{3}+u_{4}}{h^2} + Q_i = 0
$$

Note that in the above equations for $i=2$ and $i=3$, $u_2$ appears in both equations. This suggests that we will have to solve the system of equations simultaneously for all values of $u_i$. The way that we do this is to construct a matrix holding the coefficients for each equation. We'll address the first and last rows of this matrix later, but the entries that in the 2nd-through (nx-1)-th rows look like this:
$$\underline{\underline{D}}=
\begin{pmatrix}
D_{0,0} & D_{0,1} & 0 & 0 & \cdots & 0 &0  & 0\\
k/(\Delta x)^2 & -2k/(\Delta x)^2 & k/(\Delta x)^2  &0  &\cdots & 0 & 0 & 0\\
0 & k/(\Delta x)^2 & -2k/(\Delta x)^2 & k/(\Delta x)^2   &\cdots & 0 & 0 & 0\\
\vdots & \vdots & \vdots & \vdots & \ddots& \vdots & \vdots & \vdots \\
0 & 0 & 0 & 0 & \cdots & k/(\Delta x)^2 & -2k/(\Delta x)^2 & k/(\Delta x)^2 \\
0 & 0 & 0 & 0 & \cdots & 0 & D_{nx-1,nx-2} & D_{nx-1,nx}
\end{pmatrix}
$$
Note that if we factor out the term $1/(\Delta x)^2$ from the matrix $\underline{\underline{D}}$, we are left with:
$$
\underline{\underline{D}}=\frac{k}{(\Delta x)^2}
\begin{pmatrix}
\frac{\Delta x^2}{k}D_{0,0} & \frac{\Delta x^2}{k}D_{0,1} & 0 & 0 & \cdots & 0 &0  & 0\\
1 & -2 & 1  &0  &\cdots & 0 & 0 & 0\\
0 & 1 & -2 & 1   &\cdots & 0 & 0 & 0\\
\vdots & \vdots & \vdots & \vdots & \ddots& \vdots & \vdots & \vdots \\
0 & 0 & 0 & 0 & \cdots & 1 & -2 & 1 \\
0 & 0 & 0 & 0 & \cdots & 0 & \frac{\Delta x^2}{k}D_{nx-1,nx-2} & \frac{\Delta x^2}{k}D_{nx-1,nx-1}
\end{pmatrix} %= \frac{k}{(\Delta x)^2} \underline{\underline{L}}(nx)
$$

We will solve the steady-state diffusion of heat with a source term:
$$
k \frac{\partial^2 T}{\partial x^2} + H = 0
$$
Subject to the boundary conditions
$$
\begin{split}
 T\bigg|_{x=0 \mathrm{ m}} = 288 \mathrm{K}\\
 T\bigg|_{x=100 \mathrm{ m}} = 295 \mathrm{K}\\
 \end{split}
$$
on the domain
$$
0 \mathrm{ m} \le x \le 100 \mathrm{ m}
$$

We will use the value $k=2$ W/m/K and $H = 0.01$ W/m$^3$

**Question 1** What is the solution to this equation subject to the boundary conditions listed above? (show your work)

<!-- grade_id: response1, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**Question 2:** Define a variable representing the number of grid points `nx` at which we would like to obtain a solution. 
1. For now you can set nx=10. 
2. Create a vector `x` containing the values of $\underline{x}$ at each grid point.
3. Define the grid spacing `dx`.
4. Create a column vector of zeros with dimensions $nx\times 1$ called `R` (right hand side).
5. Define the constants H (internal heating) and k (thermal conductivity)


<!-- grade_id: sol1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test1, points: 2, grade: true, solution: false -->
```python
assert(R.shape == (nx,1))
assert(nx == 10)
assert(dx == x[1]-x[0])
assert(x[0]==0.0)
assert(x[-1]==100.0)
```

**Question 3:** List the non-zero entries $D_{i,j}$ for row $nx-2$ of the matrix $\underline{\underline{D}}$:

<!-- grade_id: sol2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Question 5** What is the significance of the first row of D and the last row of D? Hint: What is special about the first and last values $T_0$ and $T_{nx-1}$?

<!-- grade_id: sol3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Question 6** What will the first and last rows of the left-hand-side matrix $\underline{\underline{D}}$ and the right hand side vector $\underline{R}$ look like? Describe very clearly in words and/or use markdown.

<!-- grade_id: sol4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Question 7** Add code to construct the diffusion matrix $\underline{\underline{D}}$. The easiest way to do this is to first use a `for' loop to populate rows 1 through (nx-2) with the appropriate entries, as described in class. Then, add a couple of lines of code to change the first and last rows of D to enforce the boundary conditions as you worked out in the previous step. To implement the boundary conditions, you may also need to change the first and/or last entries in the vector $\underline{R}$.

<!-- grade_id: sol5, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test5, points: 3, grade: true, solution: false -->
```python
assert( R[0] == 288.0 )
assert( R[-1] == 295.0 )
assert( (D[0][1:]==0).all() )
assert( (D[-1][:-1]==0).all() )
```

**Question 8** Obtain a solution for T by solving the linear system $\underline{\underline{D}}\ \underline{T}=\underline{R}$. Hint: remember that there routine in `numpy.linalg` for solving linear systems.

<!-- grade_id: sol6, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test6, points: 4, grade: true, solution: false -->
```python
assert( np.isclose(T[0] ,288.0) )
assert( np.isclose(T[-1],295.0) )
assert( np.isclose(T[5],298.0617284) )
```

**Question 9** Make a plot showing x on the horizontal axis, and your numerical solution T as well as the analytic solution that you worked out above. Make sure to choose line types and colors so that both solutions can be seen and include a legend and axis labels with units.

<!-- grade_id: sol7, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Question 10** Write some code to calculate and display the heat flux (in W/m^2) at the left boundary (x=0) and at the right boundary (x=100 m) by using your numerical solution and a finite difference approximation to the derivative of temperature with respect to x.

<!-- grade_id: sol8, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Question 11** Use Fourier's law and your analytic solution to calculate the heat flux at the boundaries. How close is your numerical approximation?

<!-- grade_id: sol9, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE