# Lab Final Problem 1

Write a python function (not a script) called ‘myfact’ to calculate n! (the factorial of n). As a reminder, the factorial of an integer n is defined as:

i.e., n! is the product of all of the integers from 1 to n.

You must use a ‘for’ or ‘while’ loop somewhere in your code to receive full credit. Your function should take a single integer argument (n) and return a single integer value (the factorial of n). Note that n is a scalar, not a vector. For instance, if I call your function as myfact(5), it should return `5*4*3*2*1 = 120`. Bonus: Make your function raise a ValueError() if n is not an integer or if n < 1.

<!-- grade_id: answer, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tests, points: 3, grade: true, solution: false -->
```python
# autograder tests
assert( myfact(1) == 1)
assert( myfact(5) == 120)
assert( myfact(10) == 3628800)
```