# Introduction

This week we will be studying diffusion in the context of using morphological features to date fault scarps. Given a field example of a fault scarp, what can we say about the time that has elapsed since the occurrence of the earthquake that produced the offset? The scenario is illustrated in Figure 1.

![alt text](Scarp_Cartoons.png)
**Figure 1:** (top) Evolution of surface topography above a fault showing topography before, immediately after, and quite a while after a large earthquake. (bottom) Initial condition used in our modeling exercise.

As we discussed in class (and in your textbook), transport processes on hill slopes are described by a first-order rate law that links the volume flux of soil to the gradient of the surface:

$$
q=-\kappa \frac{\partial y}{\partial x}
$$
where $y$ is the elevation of the land surface above some datum and x is horizontal position (Figure 1). The diffusivity $\kappa$ is a site-specific constant that can be determined through modeling of topography associated with scarps of known age. The first-order rate law above, combined with mass balance considerations, leads to a diffusion equation that describes how surface morphology changes over time:

$$
\frac{\partial y}{\partial t} = \kappa \frac{\partial^2 y}{\partial x^2}
$$

We will solve this equation numerically by approximating the time derivative on the left hand side using the Forward Euler method as we did in previous labs and the spatial derivative on the right hand side using the same method used last week. The overall process is:


1. Write down the complete set of equations, initial, and boundary conditions that you will solve.
1. Write down the finite difference approximation to the equations that you will solve.
1. Define the time step $\Delta t$ and the spatial grid defined by the vector $\underline{x}$ and the grid spacing $\Delta x$ to be used.
1. Populate a vector $\underline{y}$ with the initial condition. This is the solution to the problem (i.e., values of the dependent variable) at time $t^1=t_0$, i.e. for $n=1$. Call this solution `yold`. Define a scalar called `time` to store the total elapsed time and set it equal to 0.
1. Obtain a solution for $\underline{y}$ at timestep $n+1$ ($\underline{y}^{n+1}$). Call this solution `ynew`.
1. Save or plot the output if necessary (usually it is not possible to save all of the output because of the limited amount of RAM available for storage of solutions).
1. Increase n by 1 (advance to the next time step). Make the assignment yold=ynew, since the solution that we just obtained is now the *old* solution. Add $\Delta t$ to `time`.
1. Repeat steps 5-7 as necessary until the elapsed simulation time reaches the desired total elapsed time.

Recall that last week we used subscripts of the dependent variable to denote spatial indices and superscripts to denote temporal indices. In this week's exercise, subscripts refer to spatial indices and superscripts refer to temporal indices (time step number). For instance, 
$$
y_i^n
$$
indicates the value of the solution $y$ at the $n$-th time step ($t=t^n$) and at the $i$-th grid point ($x=x_i$). Recall that for the Forward Euler method, we approximate an equation that relates the time derivative of a function $y$ to the value of some function of space and time $f(x,y)$:
$$
\frac{\partial y}{\partial t} = f(y,t)
$$
as
$$
\frac{y_i^{n}-y_i^{n-1}}{\Delta t} = f(y^{n-1},t^{n-1})
$$
Note that the right hand side of the above equation includes a non-specific function $f$. For the problem that we're trying to solve,
$$
f(y,t) = \kappa \frac{\partial^2 y}{\partial x^2}
$$
We can then write the Forward Euler approximation to this equation as:
$$
\frac{y_i^{n}-y_i^{n-1}}{\Delta t} = \kappa \left( \left.\frac{\partial^2 y}{\partial x^2}\right|_{\substack{t=t^{n-1}\\x=x_{i\ \ }}} \right)
$$
Substituting in the expression that we used last week to approximate the second derivative, we have:
$$
\left.\frac{\partial^2 y}{\partial x^2}\right|_{\substack{t=t^{n-1}\\x=x_{i\ \ }}} = \frac{1}{\Delta x^2}\left(y_{i-1}^{n-1} -2 y_i^{n-1} + y_{i+1}^{n-1} \right)
$$
Now, we can put together previous equations to obtain:
$$
\frac{y_i^{n}-y_i^{n-1}}{\Delta t} = \kappa\frac{1}{\Delta x^2}\left(y_{i-1}^{n-1} -2 y_i^{n-1} + y_{i+1}^{n-1} \right)
$$
Note that the above expression can be re-arranged to give you an expression for each $y_i$ at timestep $n$ in terms of only quantities defined at timestep $n-1$.

# Statement of Problem

You will create a MATLAB script to solve a 1-dimensional diffusion problem in both space and time. This week's exercises assemble what you did previously to solve an initial value problem with what you did last week to solve a boundary value problem. The problem that you will solve is:
$$
\frac{\partial y}{\partial t} = \kappa \frac{\partial^2 y}{\partial x^2}
$$
Subject to the initial condition:
$$
y(x,0) = \left\{
     \begin{array}{lr}
       2 & : x > 0\\
       0 & : x \le 0
            \end{array}
   \right.
$$
and boundary conditions:
$$
\begin{split}
y(-150,t) = 0\\
y(150,t) = 2
\end{split}
$$
Thus, the domain in which you need to calculate a solution is $x\in[ -150 m, +150 m]$. You will start at time $t=0$ and take time steps of size $\Delta t$. The value of $\kappa$ for this problem is 1.5 m$^2$/kyr (kyr means 1000 years).

**0.** What is the solution to the slope diffusion equation with boundary conditions specified above in the limit as $t\to \infty$? *Hint: The solution as time goes to infinity is the same as the steady-state solution* **Show your work.**

<!-- grade_id: ans0, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**1**. Start by writing down the finite-difference approximation to the diffusion equation using the Forward Euler scheme for time integration. Solve the expression for $y_i^{n}$. You should have an equation that has only $y_i^n$ on the left hand side and $y_{i-1}^{n-1}$, $y_i^{n-1}$, and $y_{i+1}^{n-1}$ on the right hand side. Note that all of these right hand side terms involve values of $\underline{y}$ at time $n-1$.

<!-- grade_id: ans1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

The finite difference scheme above can be expressed in the form of $\underline{y}^n = \underline{\underline{M}}\ \underline{y}^{n-1}$ (i.e. the solution at time step $n$ can be obtained by multiplying a yet-to-be-determined matrix with the vector containing the solution at time $n-1$. Write down expressions for the entries in the matrix $\underline{\underline{M}}$. Hint: start with what you did in the previous step and write out the expressions for $y_i^n$ for spatial indices $i=2$ and $i=3$. Do you recognize a pattern? You may be able to express the matrix $\underline{\underline{M }}$ in terms of the diffusion matrix $\underline{\underline{D}}$ from last week and the identity matrix (except for the first and last rows).

<!-- grade_id: ans1-5, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

What do you need to do to ensure that the boundary conditions are satisfied at each timestep? Hint: they are enforced as part of the initial condition. You need to think about how to maintain their values at subsequent time steps.

<!-- grade_id: ans2, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Add code to:
1. create a uniformly-spaced x-vector (as you did last week)
2. define `tmin` and `tmax` (tmax should be set to 10,000 years)
2. crate a uniformly spaced t-vector (as you did earlier this quarter)
3. and define the grid spacing dx.
4. define the time step dt. For now, use 201 grid points in the x-direction. Choose a time step $dt=\frac{(\Delta x)^2}{4\kappa}$.

**Don't forget to check for consistent dimensions throughout your code. The autograder tests assume that you work in meters and kyr for distance and time.**

<!-- grade_id: ans3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: grade1, points: 5, grade: true, solution: false -->
```python
assert(x[0]==-150)
assert(x[-1]==150)
assert(kappa==1.5)
assert(dt == 0.375)
assert(dx == 1.50)
assert(t[0] == 0.)
assert( np.isclose(t[-1],1e1) )
```

Create a vector `y0` to hold the initial condition (values of y at each grid point at time 0):

<!-- grade_id: cell6-1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell6-2, points: 2, grade: true, solution: false -->
```python
assert( y0.shape == (nx,1) )
assert( np.all( y0[x<=0] == 0.0 ) )
assert( np.all( y0[x>0] == 2.0 ) )
```

Add code to your m-file to construct the matrix $\underline{\underline{M}}$ that you defined above.

<!-- grade_id: cell6-5, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Add code to your m-file to solve the time-dependent diffusion equation starting at $t=0$ and ending at $t=tmax$ using the procedure outlined in the introduction.

<!-- grade_id: cell8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Produce a plot showing the scarp profiles predicted at 0 years (initial condition), and at 1000 year intervals until t=10,000 years. Hint: You should be able to come up with a method to save and/or plot these results within the loop that you use to do your time stepping calculation in the previous part.

<!-- grade_id: plot, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

The file `scarp_data.txt` contains a dataset with vectors xdata and ydata that contain measurements of the surface elevation along a transect of an alluvial fan with a fault scarp produced by a holocene earthquake. The scarp is centered at $x=0$m. Determine the age of earthquake that produced the offset and estimate the uncertainty associated with your model age. Make a plot showing the data (you can use the `scatter` command) and the modeled present-day elevation profile for your preferred scarp age. Make sure that it is appropriately labeled and annotated.

```python
# THIS BLOCK READS AND PLOTS THE SCARP DATA SO THAT YOU CAN SEE WHAT IT LOOKS LIKE

import numpy as np
import matplotlib.pyplot as plt
scarp_x, scarp_y  = np.loadtxt('scarp_data.txt',usecols=(0,1),unpack=True)
plt.plot(scarp_x,scarp_y,'r.')
plt.xlabel('x-postion (m)')
plt.ylabel('elevation (m)')
plt.show()
```

<!-- grade_id: age, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

 Try modifying your code so that the time step is larger - try increasing it by a factor of 4: 
 $dt=\frac{(\Delta x)^2}{\kappa}$. What happens? Produce a plot below showing your calculated solution at t=10,000 years using this larger time step size. It will probably be easiest to just copy-and-paste the code to set up and solve the problem here, and modify only the time step. (Use the two cells below for code and markdown to address this question.)

<!-- grade_id: cell13, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell14, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE